```python
def print_inverse(n, base=10):
    remainder   = 1
    denominator = n
    for _ in range(10):
        numerator = remainder * base
        quotient, remainder = divmod(numerator, denominator)
        print(quotient, end=", ")
    print()
```

```python
print_inverse(7)
```

```python
type(print_inverse(7))
```

課題-2) 課題-1で定義した `print_inverse()` は、組み込みの `print()` 関数を遣って画面出力をしており、関数の戻り値を返していない。次の例を参考にして、数列 (`int`値を要素とする`list`) を戻り値とするように改変せよ。

```
def test():
    l = []
    for i in range(10):
        l.append(i)
    return l
```

```python
def test():
    l = []
    for i in range(10):
        l.append(i)
    return l

test()
```

<!-- grade_id: cell-9a90b3445b15156e, points: 1, grade: true, solution: true -->
```python
def inverse(n, base=10):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
type(inverse(4))
```

```python
inverse(4)
```

```python
inverse(7)
```

```python

```

---
以下、採点用のセルにつき編集しないようにして下さい:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-b698c617121292c0, points: 1, grade: true, solution: false -->
```python
assert_true(isinstance(inverse(4), list))
```

<!-- grade_id: cell-19f3ab02a131a394, points: 1, grade: true, solution: false -->
```python
assert_equal(
    list(inverse(4)),
    [2, 5, 0, 0, 0, 0, 0, 0, 0, 0])
```

<!-- grade_id: cell-9cda36647da9df56, points: 1, grade: true, solution: false -->
```python
assert_equal(
    list(inverse(7)), 
    [1, 4, 2, 8, 5, 7, 1, 4, 2, 8])
```

<!-- grade_id: cell-3e385a01c8ce864f, points: 1, grade: true, solution: false -->
```python
assert_equal(
    list(inverse(3,60)),
    [20, 0, 0, 0, 0, 0, 0, 0, 0, 0])
```