```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# Assignment (4)

次の `inverse()` は、有限小数の末尾に 0 を出力してループを停止するように改良したものである。さらに、循環小数は、循環を検出してループを停止している。その場合は末尾に 0 以外の数が出力される:

```python
def inverse(n, base=10):
    denominator = n
    remainder   = 1 # initial remainder
    quotient    = 0 # dummy
    repeat = set()
    l = []
    while not quotient in repeat:
        repeat.add(quotient)
        numerator = remainder * base
        quotient, remainder = divmod(numerator, denominator)
        l.append(quotient)
        if remainder == 0:
            l.append(0)
            break
    return l
```

```python
inverse(4)
```

```python
(inverse(4))[-1]
```

```python
inverse(7)
```

```python
(inverse(7))[-1]
```

```python
[(inverse(n))[-1] for n in range(2,10)]
```

```python
[0, 3, 0, 0, 6, 1, 0, 1].count(0)
```

```python
[(inverse(n,60))[-1] for n in range(2,60)]
```

```python
[(inverse(n,60))[-1] for n in range(2,60)].count(0)
```

課題-4) 上の関数 `inverse()` は、有限小数の末尾に `0` を出力してループを停止するように改良したものである。さらに、循環小数は、循環を検出してループを停止している。その場合は末尾に `0` 以外の数が出力される。

課題-4.1) 10進数の有限小数で $\frac{1}{n} \ (1 < n < 10)$ で表される8個の数の内、有限小数を数えよ。

課題-4.2) 30進数の有限小数で $\frac{1}{n} \ (1 < n < 30)$ で表される28個の数の内、有限小数を数えよ。

課題-4.3) 60進数の有限小数で $\frac{1}{n} \ (1 < n < 60)$ で表される58個の数の内、有限小数を数えよ。

課題-6.4) 360進数の有限小数で $\frac{1}{n} \ (1 < n < 360)$ で表される358個の数の内、有限小数を数えよ。

<!-- grade_id: cell-572674e6b04017ee, points: 1, grade: true, solution: true -->
```python
def count_finite_decimal(base=10):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

<!-- grade_id: cell-91377f1eb231ac37, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(count_finite_decimal(10),4)
```

<!-- grade_id: cell-deb94c1a9aea39e6, points: 1, grade: true, solution: false -->
```python
assert_equal(count_finite_decimal(30),16)
```

<!-- grade_id: cell-4eb6c6d19766c45a, points: 1, grade: true, solution: false -->
```python
assert_equal(count_finite_decimal(60),24)
```

<!-- grade_id: cell-438037b7d72f0579, points: 1, grade: true, solution: false -->
```python
assert_equal(count_finite_decimal(360),56)
```