```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# Assignment (2)

```python
def print_inverse(n, base=10):
    remainder   = 1
    denominator = n
    for _ in range(10):
        numerator = remainder * base
        quotient, remainder = divmod(numerator, denominator)
        print(quotient, end=", ")
    print()
```

```python
print_inverse(7)
```

```python
type(print_inverse(7))
```

```python
def test():
    l = []
    for i in range(10):
        l.append(i)
    return l
```

```python
test()
```

```python
type(test())
```

課題-2) 課題-1で定義した `print_inverse()` は、組み込みの `print()` 関数を遣って画面出力をしており、関数の戻り値を返していない (`None`)。次の例を参考にして、数列 (`int`値を要素とする`list`) を戻り値とするように6行目と7行目の二つの `print()` を置き換えて改変せよ。

```
def test():
    l = []
    for i in range(10):
        l.append(i)
    return l
```

<!-- grade_id: cell-9a90b3445b15156e, points: 0, grade: false, solution: true -->
```python
def inverse(n, base=10):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
type(inverse(4))
```

```python
inverse(4)
```

```python
inverse(7)
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-b698c617121292c0, points: 1, grade: true, solution: false -->
```python
assert_true(isinstance(inverse(4), list))
```

<!-- grade_id: cell-19f3ab02a131a394, points: 1, grade: true, solution: false -->
```python
assert_equal(
    inverse(4),
    [2, 5, 0, 0, 0, 0, 0, 0, 0, 0])
```

<!-- grade_id: cell-9cda36647da9df56, points: 1, grade: true, solution: false -->
```python
assert_equal(
    inverse(7), 
    [1, 4, 2, 8, 5, 7, 1, 4, 2, 8])
```

<!-- grade_id: cell-3e385a01c8ce864f, points: 1, grade: true, solution: false -->
```python
assert_equal(
    inverse(3,60),
    [20, 0, 0, 0, 0, 0, 0, 0, 0, 0])
```