```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# Assignment (1)

```python
base = 10
remainder   = 1
denominator = 7
for _ in range(10):
    numerator = remainder * base
    quotient, remainder = divmod(numerator, denominator)
    print(quotient, end=", ")
print()
```

課題-1) 上のセルをボディとする関数 `print_inverse()` を定義せよ。
* 第1引数 `n` (逆数を求める数) と第2引数 `base` (基数) を持つ
* 第2引数の `base` は省略可能で、省略時の値は `10` (10進数) とする
* 上のセルの内容 (第1行以外の第2行から第8行まで) を下のセルの関数ボディ (コメント行 ```# YOUR CODE HERE``` 以下) へコピペし、適切にインデントする
* コメント行 ```# YOUR CODE HERE``` は削除してもしなくてもよい
* コメント行の下の ```raise NotImplementedError()``` は削除し、本文と置き換える
* 上のセルの第1行 `base = 10` は関数の省略可能な引数として定義しているのでコピペしない
* ボディでは、関数の第1引数 `n` を `denominator` に代入する
* `for` ループの繰り返し回数は、何桁まで計算するかであり、基数 `base` とは関係ないことに注意

<!-- grade_id: cell-64114849c9fbcae5, points: 1, grade: true, solution: true -->
```python
def print_inverse(n, base=10):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
type(print_inverse(4))
```

```python
print_inverse(4)
```

```python
print_inverse(7)
```

```python

```

```python

```

```python

```

```python
for n in range(2,10):
    print_inverse(n)
```

```python
for n in range(2,10):
    print(1/n)
```

---
以下、採点用のセルにつき編集できない:

<!-- grade_id: cell-08e89639eaba656e, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-9585afce246c2e75, points: 0, grade: true, solution: false -->
```python
assert_true(isinstance(print_inverse(4),type(None)))
```