```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# 近代数学とプログラミング (5)

### 二進法とブール代数
* [Boolean algebra - Wikipedia](https://en.wikipedia.org/wiki/Boolean_algebra)

```python
import itertools

list(itertools.product([0,1],repeat=2))
```

```python
for i,j in itertools.product([0,1],repeat=2):
    print(i&j, bool(i&j), bool(i)&bool(j))
```

```python
for i,j in itertools.product([0,1],repeat=2):
    print(i|j, bool(i|j), bool(i)|bool(j))
```

```python
for i,j in itertools.product([0,1],repeat=2):
    print(i^j, bool(i^j), bool(i)^bool(j))
```

```python

```

```python

```

```python

```

### ブール演算と四則演算

```python
import itertools

for i,j in itertools.product([0,1],repeat=2):
    print(i&j, i*j, '{:02b}'.format(i*j))
```

```python
for i,j in itertools.product([0,1],repeat=2):
    print(i|j, i+j, '{:02b}'.format(i+j))
```

```python
for i,j in itertools.product([0,1],repeat=2):
    print(i^j, i+j, '{:02b}'.format(i+j))
```

```python

```

```python

```

```python

```

和 (or) は、四則演算や直感と対応が難しい

```python
for i,j in itertools.product([0,1],repeat=2):
    print(i|j, i+j, '{:02b}'.format(i+j), divmod(i+j,2), sum(divmod(i+j,2)), ((i+j)%2)+(i+j)//2)
```

```python
for i,j in itertools.product([0,1],repeat=2):
    print(i^j, i+j, '{:02b}'.format(i+j), divmod(i+j,2), divmod(i+j,2)[1], (i+j)%2)
```

```python

```

```python

```

```python

```

課題-5) 一桁の二進数の和を求める関数 ``binary_add()`` を定義せよ。戻り値は整数 $0$, $1$ のリストで表すこと。

<!-- grade_id: cell-3172796abe850fcd, points: 0, grade: false, solution: true -->
```python
def binary_add(a, b):
    if not (a >=0 and b >= 0 and a < 2 and b < 2):
        raise ValueError
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
binary_add(0, 1)
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-edbe9b33d9963ab9, points: 1, grade: true, solution: false -->
```python
assert_equal(binary_add(0, 1), [0, 1])
```