```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# Assingment

* [Addition - Wikipedia](https://en.wikipedia.org/wiki/Addition)

Step-1) 与えられた整数の桁を逆に並べたリストを作る

```python
str(1024)
```

```python
list(reversed(str(1024)))
```

```python
list(map(int,reversed(str(1024))))
```

```python

```

```python

```

```python

```

Step-2) 与えられた２つの整数 $a$, $b$ の下桁から上桁までの順に並べる

```python
a = list(map(int,reversed(str(256))))
b = list(map(int,reversed(str(64))))
a, b
```

```python
from itertools import zip_longest
```

```python
for n, m in zip_longest(a, b, fillvalue=0):
    print(n, m)
```

```python
for n, m in zip_longest(a, b, fillvalue=0):
    print(n+m)
```

```python
for n, m in zip_longest(a, b, fillvalue=0):
    print(divmod(n+m, 10))
```

```python

```

```python

```

```python

```

Step-3) 繰り上がりを処理する

```python
q=0
for n, m in zip_longest(a, b, fillvalue=0):
    q,r = divmod(q+n+m, 10)
    print(r)
print(q)
```

```python

```

```python

```

```python

```

Step-4) 戻り値を準備する

```python
l=[]
q=0
for n, m in zip_longest(a, b, fillvalue=0):
    q,r = divmod(q+n+m, 10)
    l.append(r)
l.append(q)
l
```

```python

```

```python

```

```python

```

Step-4) 逆順にして10進法の位取り記法の整数に戻す

```python
l=[]
q=0
for n, m in zip_longest(a, b, fillvalue=0):
    q,r = divmod(q+n+m, 10)
    l.append(r)
l.append(q)
l
```

```python
list(reversed(l))
```

```python
list(map(str,(reversed(l))))
```

```python
''.join(map(str,(reversed(l))))
```

```python
int(''.join(map(str,(reversed(l)))))
```

```python

```

```python

```

```python

```

課題) 10進法の位取り記法で表される数を桁毎に計算するプログラムを作成せよ。

<!-- grade_id: cell-4a12ad3a538e7ec2, points: 0, grade: false, solution: true -->
```python
from itertools import zip_longest

def long_addition(a, b, base=10):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
long_addition(12345, 6789), 12345+6789
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-079d68b98f2d72b7, points: 1, grade: true, solution: false -->
```python
assert_equal(long_addition(345, 456), 345+456)
```

<!-- grade_id: cell-71f69003f514f4d8, points: 1, grade: true, solution: false -->
```python
assert_equal(long_addition(12345, 6789), 12345+6789)
```