```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# Assignment (3)

```python
def inverse(n, base=10):
    remainder   = 1
    denominator = n
    l = []
    for _ in range(10):
        numerator = remainder * base
        quotient, remainder = divmod(numerator, denominator)
        l.append(quotient)
    return l
```

```python
inverse(4)
```

Output:
```
[2, 5, 0, 0, 0, 0, 0, 0, 0, 0]
```

```python
def test():
    for i in range(10):
        if i > 2:
            break
        print(i)
```

```python
test()
```

Output:
```
0
1
2
```

課題-3) ループは `break` で中断できる。`inverse()` 関数の `for` ループ内で求めている `remainder` が `0` の時、ループを中断するように改変せよ。
中断に際して、有限小数であることを判定可能とするため、リストの末尾に `0` が付加されるように配慮すること。

<!-- grade_id: cell-3b430a715dc0a8aa, points: 0, grade: false, solution: true -->
```python
def inverse(n, base=10):
    ### BEGIN SOLUTION
    remainder   = 1
    denominator = n
    l = []
    for _ in range(10):
        numerator = remainder * base
        quotient, remainder = divmod(numerator, denominator)
        l.append(quotient)
        if remainder == 0:
            l.append(0)
            break
    return l
    ### END SOLUTION
```

```python
inverse(4)
```

Output:
```
[2, 5, 0]
```

```python
(inverse(4))[-1]
```

Output:
```
0
```

```python
for i in range(2,10):
    print((inverse(i))[-1])
```

Output:
```
0
3
0
0
6
8
0
1
```

```python

```

```python

```

```python

```

## 有限小数を判定する

```python
for i in range(2,10):
    if (inverse(i))[-1] == 0:
        print("1/"+str(i)+" is finite decimal")
```

Output:
```
1/2 is finite decimal
1/4 is finite decimal
1/5 is finite decimal
1/8 is finite decimal
```

```python
for i in range(2,60):
    if (inverse(i,60))[-1] == 0:
        print("1/"+str(i)+" is finite decimal")
```

Output:
```
1/2 is finite decimal
1/3 is finite decimal
1/4 is finite decimal
1/5 is finite decimal
1/6 is finite decimal
1/8 is finite decimal
1/9 is finite decimal
1/10 is finite decimal
1/12 is finite decimal
1/15 is finite decimal
1/16 is finite decimal
1/18 is finite decimal
1/20 is finite decimal
1/24 is finite decimal
1/25 is finite decimal
1/27 is finite decimal
1/30 is finite decimal
1/32 is finite decimal
1/36 is finite decimal
1/40 is finite decimal
1/45 is finite decimal
1/48 is finite decimal
1/50 is finite decimal
1/54 is finite decimal
```

```python
inverse(3)
```

Output:
```
[3, 3, 3, 3, 3, 3, 3, 3, 3, 3]
```

```python
inverse(3,60)
```

Output:
```
[20, 0]
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-942b60b4f3cf0d30, points: 1, grade: true, solution: false -->
```python
assert_equal(
    inverse(4),
    [2, 5, 0])
```

<!-- grade_id: cell-5f7ac7348e0f01fb, points: 1, grade: true, solution: false -->
```python
assert_equal(
    (inverse(4))[-1],
    0)
```

<!-- grade_id: cell-d4fefd32d24c096c, points: 1, grade: true, solution: false -->
```python
assert_equal(
    inverse(7), 
    [1, 4, 2, 8, 5, 7, 1, 4, 2, 8])
```