```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# 近代数学とプログラミング (6)

## set comprehension

* [List comprehension - Wikipedia](https://en.wikipedia.org/wiki/List_comprehension#Python)
  - [Set-builder notation - Wikipedia](https://en.wikipedia.org/wiki/Set-builder_notation)

$$
  S = \left\{ n \in \mathbb{Z} \vert n \not\equiv 0 \ \mathrm{mod}\ 2\right\}
$$

```python
o = {n for n in range(10) if n % 2}
o
```

```python
type(o)
```

```python
e = {n for n in range(10) if not n % 2}
e
```

```python
e&o, e|o, e^o
```

```python
e-o, o-e
```

```python

```

```python

```

```python

```

## list comprehension

```python
o = [n for n in range(10) if n % 2]
o
```

```python
type(o)
```

```python
e = [n for n in range(10) if not n % 2]
e
```

```python
o&e, o|e, o^e
```

```python
o+e
```

```python

```

```python

```

課題-1)

<!-- grade_id: cell-a076cfa8adba5833, points: 0, grade: false, solution: true -->
```python
def odd_numbers(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
odd_numbers(10)
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-c158e13c923b33c2, points: 1, grade: true, solution: false -->
```python
assert_equal(odd_numbers(10), [1, 3, 5, 7, 9])
```