```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# 課題-2) Plimpton 322 に記載された数表を生成する

```python
import math
import itertools
```

```python
def generate_p(i_max=2, j_max=2, k_max=2):
    l = []
    for i in range(i_max):
        for j in range(j_max):
            for k in range(k_max):
                l.append(2**i * 3**j * 5**k)
    return sorted(l)

generate_p()
```

```python
pp = itertools.product(generate_p(), generate_p())

list(pp);
```

```python
pp = itertools.product(generate_p(), repeat=2)

list(pp);
```

```python
pq = []
for p, q in itertools.product(generate_p(), repeat=2):
    if p > q and math.gcd(p,q) ==1:
        if p/q >= 9/5 and p/q <= 12/5:
            pq.append((p,q))

pq
```

```python
pq = []
for p, q in itertools.product(generate_p(4,3,2), repeat=2):
    if p > q and math.gcd(p,q) ==1:
        if p/q >= 9/5 and p/q <= 12/5:
            pq.append((p,q))

pq
```

```python
sorted(pq, reverse=True, key=lambda x: x[0]/x[1])
```

```python

```

```python

```

```python

```

課題-2) Plimton 322の数表を求める関数 `plimpton322()` を定義せよ:
* 引数 i_max, j_max, k_max を取り、デフォルト値はそれぞれ 2 とする
* `generate_p(i_max, j_max, k_max)` の直積の要素 `p, q` のうち、次の条件を満たすものを抜き出す
  - `p > q`
  - `math.gcd(p,q) == 1` (互いに素)
  - `p/q >= 9/5`
  - `p/q <= 12/5`
* 戻り値は、`p/q` をキーとして降順で返す

<!-- grade_id: cell-b54242a64689bdeb, points: 0, grade: false, solution: true -->
```python
def plimpton322(i_max=2, j_max=2, k_max=2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plimpton322()
```

```python
plimpton322(4,3,2)
```

```python

```

```python

```

```python

```

```python
for p, q in plimpton322(4,3,2):
    a = 2 * p * q
    b = p**2 - q**2
    d = p**2 + q**2
    print((d/a)**2, b, d)
```

```python
import pandas as pd
import requests

res = requests.get("https://en.wikipedia.org/wiki/Plimpton_322").text

df2=pd.read_html(res)[1]
df2
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-b4f212f1a4f0d324, points: 1, grade: true, solution: false -->
```python
assert_equal(plimpton322(), [(2, 1)])
```

<!-- grade_id: cell-4fdfe4778425eec3, points: 1, grade: true, solution: false -->
```python
assert_equal(plimpton322(4,3,2), [(12, 5), (9, 4), (20, 9), (2, 1), (15, 8), (9, 5)])
```