```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# 近代数学とプログラミング

## 関数論と四則演算

* (1903 &ndash; 1995) [Alonzo Church - Wikipedia](https://en.wikipedia.org/wiki/Alonzo_Church)
* [Church encoding - Wikipedia](https://en.wikipedia.org/wiki/Church_encoding)
  - [Adding church numerals using lambda functions in python - Stack Overflow](https://stackoverflow.com/questions/24905241/adding-church-numerals-using-lambda-functions-in-python)

```python
def f(x):
    return x+1
```

```python
0, f(0), f(f(0)), f(f(f(0)))
```

Output:
```
(0, 1, 2, 3)
```

```python
f = lambda x: x+1
```

```python
0, f(0), f(f(0)), f(f(f(0)))
```

Output:
```
(0, 1, 2, 3)
```

```python

```

```python

```

```python

```

### 関数の合成回数で序数を定義する
* [Function composition - Wikipedia](https://en.wikipedia.org/wiki/Function_composition)

```python
def zero(f):
    return lambda x: x
```

```python
def one(f):
    return lambda x: f(x)
```

```python
def two(f):
    return lambda x: f(f(x))
```

```python
def three(f):
    return lambda x: f(f(f(x)))
```

```python
f
```

Output:
```
<function __main__.<lambda>(x)>
```

```python
zero
```

Output:
```
<function __main__.zero(f)>
```

```python
zero(f)
```

Output:
```
<function __main__.zero.<locals>.<lambda>(x)>
```

```python
(zero(f))(0)
```

Output:
```
0
```

```python
(one(f))(0)
```

Output:
```
1
```

```python
(two(f))(0)
```

Output:
```
2
```

```python
(three(f))(0)
```

Output:
```
3
```

```python

```

```python

```

```python

```

### 関数の加算

```python
def church_add(m, n):
    return lambda f: lambda x: m(f)(n(f)(x))
```

```python
church_add(one,two)
```

Output:
```
<function __main__.church_add.<locals>.<lambda>(f)>
```

```python
church_add(one,two)(f)
```

Output:
```
<function __main__.church_add.<locals>.<lambda>.<locals>.<lambda>(x)>
```

```python
(church_add(one,two)(f))(0)
```

Output:
```
3
```

```python
(church_add(three,two)(f))(0)
```

Output:
```
5
```

```python

```

```python

```

```python

```

### 関数の乗算

```python
def church_mul(m, n):
    return lambda f: lambda x: m(n(f))(x)
```

```python
church_mul(one, two)(f)(0)
```

Output:
```
2
```

```python
church_mul(two, two)(f)(0)
```

Output:
```
4
```

```python
church_mul(two, three)(f)(0)
```

Output:
```
6
```

```python

```

```python

```

```python

```

### 後者関数

```python
def church_succ(n):
    return lambda f: lambda x: f(n(f)(x))
```

```python
def church_zero(f):
    return lambda x: x
```

```python
def church_one(f):
    return lambda x: (church_succ(church_zero))(f)(x)
```

```python
def church_two(f):
    return lambda x: (church_succ(church_succ(church_zero)))(f)(x)
```

```python
church_zero(f)(0)
```

Output:
```
0
```

```python
(church_succ(church_zero))(f)(0)
```

Output:
```
1
```

```python
church_one(f)(0)
```

Output:
```
1
```

```python
church_two(f)(0)
```

Output:
```
2
```

```python

```

```python

```

```python

```

課題-7) 次のように定義する ``nest_lit()`` を使って、２つの整数の足し算をする関数 ``nest_add()`` を定義せよ。

<!-- grade_id: cell-f0a3f9534a0d2fd7, points: 0, grade: false, solution: true -->
```python
def nest_list(f=lambda x: x+1, x=0):
    while True:
        yield x
        x = f(x)

def nest_add(a, b):
    ### BEGIN SOLUTION
    g = nest_list()
    {next(g) for _ in range(a)}
    {next(g) for _ in range(b)}
    return next(g)
    ### END SOLUTION
```

```python
nest_add(5,3)
```

Output:
```
8
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-a6359144b116935e, points: 1, grade: true, solution: false -->
```python
assert_equal(nest_add(5, 3), 8)
```