```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# 近代数学とプログラミング

## 集合論と序数

* [5. Data Structures — Python 3.10.5 documentation](https://docs.python.org/3/tutorial/datastructures.html#sets)
  - [Built-in Types — Python 3.10.5 documentation](https://docs.python.org/3/library/stdtypes.html#set)

### Pythonの空集合 $\phi$ の扱い:

``list``と``tuple``と``dict``と``set``

```python
[], list([]), list()
```

Output:
```
([], [], [])
```

```python
list(map(type, ([], list([]), list())))
```

Output:
```
[list, list, list]
```

```python
list(map(bool, ([], list([]), list())))
```

Output:
```
[False, False, False]
```

```python
if not []:
    print('empty')
```

Output:
```
empty
```

```python

```

```python

```

```python

```

```python
(), tuple(()), tuple()
```

Output:
```
((), (), ())
```

```python
list(map(type, ((), tuple(()), tuple())))
```

Output:
```
[tuple, tuple, tuple]
```

```python
list(map(bool, ((), tuple(()), tuple())))
```

Output:
```
[False, False, False]
```

```python
if not ():
    print('empty')
```

Output:
```
empty
```

```python

```

```python

```

```python

```

```python
{}, dict({}), dict()
```

Output:
```
({}, {}, {})
```

```python
list(map(type, ({}, dict({}), dict())))
```

Output:
```
[dict, dict, dict]
```

```python
list(map(bool, ({}, dict({}), dict())))
```

Output:
```
[False, False, False]
```

```python
if not {}:
    print('empty')
```

Output:
```
empty
```

```python
d={'jack': 4098, 'sape': 4139}
d
```

Output:
```
{'jack': 4098, 'sape': 4139}
```

```python
type(d), d['jack']
```

Output:
```
(dict, 4098)
```

```python

```

```python

```

```python

```

* [PEP 218 – Adding a Built-In Set Object Type | peps.python.org](https://peps.python.org/pep-0218/)

```python
{}, set({}), set()
```

Output:
```
({}, set(), set())
```

```python
list(map(type, ({}, set({}), set())))
```

Output:
```
[dict, set, set]
```

```python
list(map(bool, ({}, set({}), set())))
```

Output:
```
[False, False, False]
```

```python
if not set({}):
    print('empty')
```

Output:
```
empty
```

```python
s={1,2,3}
s
```

Output:
```
{1, 2, 3}
```

```python
type(s)
```

Output:
```
set
```

```python

```

```python

```

```python

```

### von Neumannの序数

* (1903&ndash;1957) [John von Neumann - Wikipedia](https://en.wikipedia.org/wiki/John_von_Neumann)
  - [Ordinal number - Wikipedia](https://en.wikipedia.org/wiki/Ordinal_number#Von_Neumann_definition_of_ordinals)
    - [Set-theoretic definition of natural numbers - Wikipedia](https://en.wikipedia.org/wiki/Set-theoretic_definition_of_natural_numbers#:~:text=Definition%20as%20von%20Neumann%20ordinals,-See%20also%3A%20Zermelo&text=The%20set%20N%20of%20natural,Peano%20axioms%20%28Goldrei%201996%29.)

ひとつ前の集合を包含する:

$$
  \begin{align}
    0 &= \{\} \\
    S(a) &= a \cup \{a\} 
  \end{align}
$$

0. $\phi$
1. $\{\phi\}$
2. $\{\phi, \{\phi\}\}$
3. $\{\phi, \{\phi, \{\phi\}\}\}$

Pythonでは、集合は集合を含むことができない。

```python
def succ(n):
    return n|{n}
```

```python
succ(set())
```

Output:
```
Input In [3], in <cell line: 1>()
----> 1 succ(set())

Input In [2], in succ(n)
      1 def succ(n):
----> 2     return n|{n}

TypeError: unhashable type: 'set'
```

```python
set()|{set()}
```

Output:
```
TypeError                                 Traceback (most recent call last)
Input In [4], in <cell line: 1>()
----> 1 set()|{set()}

TypeError: unhashable type: 'set'
```

```python
0|{0}
```

Output:
```
TypeError                                 Traceback (most recent call last)
Input In [5], in <cell line: 1>()
----> 1 0|{0}

TypeError: unsupported operand type(s) for |: 'int' and 'set'
```

```python

```

```python

```

```python

```

次のように表現する:

```python
n = [None]*10
n[0] = set()
n
```

Output:
```
[set(), None, None, None, None, None, None, None, None, None]
```

```python
n[1] = {0}|n[0]
n
```

Output:
```
[set(), {0}, None, None, None, None, None, None, None, None]
```

```python
n[2] = {1}|n[1]
n
```

Output:
```
[set(), {0}, {0, 1}, None, None, None, None, None, None, None]
```

```python
n[3] = {2}|n[2]
n
```

Output:
```
[set(), {0}, {0, 1}, {0, 1, 2}, None, None, None, None, None, None]
```

```python

```

```python

```

```python

```

```python
n = [set()]

def succ(i):
    return {i}|n[i]

for i in range(10):
    n.append(succ(i))
n
```

Output:
```
[set(),
 {0},
 {0, 1},
 {0, 1, 2},
 {0, 1, 2, 3},
 {0, 1, 2, 3, 4},
 {0, 1, 2, 3, 4, 5},
 {0, 1, 2, 3, 4, 5, 6},
 {0, 1, 2, 3, 4, 5, 6, 7},
 {0, 1, 2, 3, 4, 5, 6, 7, 8},
 {0, 1, 2, 3, 4, 5, 6, 7, 8, 9}]
```

```python
n[0], n[1]
```

Output:
```
(set(), {0})
```

```python

```

```python

```

集合の包含関係による順序が導入されているが、プログラミング的に応用は難しい･･･
* [Set (mathematics) - Wikipedia](https://en.wikipedia.org/wiki/Set_%28mathematics%29#Subsets)

$\{0,1,2\} \subset \{0,1,2,3,4\}$

```python
type(5), type(3)
```

Output:
```
(int, int)
```

```python
5<3, 5>3
```

Output:
```
(False, True)
```

```python
n[5], n[3]
```

Output:
```
({0, 1, 2, 3, 4}, {0, 1, 2})
```

```python
type(n[5]), type(n[3])
```

Output:
```
(set, set)
```

```python
n[5]<n[3], n[5]>n[3]
```

Output:
```
(False, True)
```

```python
n[5].issubset(n[3]), n[3].issubset(n[5])
```

Output:
```
(False, True)
```

```python

```

```python

```

```python

```

```python
(n[5]&n[3])==n[5], (n[5]&n[3])==n[3]
```

Output:
```
(False, True)
```

```python
n[5]&n[3], n[5]|n[3], n[5]^n[3], n[5]-n[3]
```

Output:
```
({0, 1, 2}, {0, 1, 2, 3, 4}, {3, 4}, {3, 4})
```

```python

```

```python

```

```python

```

```python
list(map(len, n))
```

Output:
```
[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
```

```python
sorted(n)
```

Output:
```
[set(),
 {0},
 {0, 1},
 {0, 1, 2},
 {0, 1, 2, 3},
 {0, 1, 2, 3, 4},
 {0, 1, 2, 3, 4, 5},
 {0, 1, 2, 3, 4, 5, 6},
 {0, 1, 2, 3, 4, 5, 6, 7},
 {0, 1, 2, 3, 4, 5, 6, 7, 8},
 {0, 1, 2, 3, 4, 5, 6, 7, 8, 9}]
```

```python
sorted(n, key=lambda x: len(x))
```

Output:
```
[set(),
 {0},
 {0, 1},
 {0, 1, 2},
 {0, 1, 2, 3},
 {0, 1, 2, 3, 4},
 {0, 1, 2, 3, 4, 5},
 {0, 1, 2, 3, 4, 5, 6},
 {0, 1, 2, 3, 4, 5, 6, 7},
 {0, 1, 2, 3, 4, 5, 6, 7, 8}]
```

```python
sorted(n, key=lambda x: -len(x))
```

Output:
```
[{0, 1, 2, 3, 4, 5, 6, 7, 8},
 {0, 1, 2, 3, 4, 5, 6, 7},
 {0, 1, 2, 3, 4, 5, 6},
 {0, 1, 2, 3, 4, 5},
 {0, 1, 2, 3, 4},
 {0, 1, 2, 3},
 {0, 1, 2},
 {0, 1},
 {0},
 set()]
```

```python

```

```python

```

```python

```

課題-2) 上で定義した自然数を表す集合のリスト``n``の要素の大小を判定する ``less_than()`` を定義せよ。

<!-- grade_id: cell-76bd4406febd5a93, points: 0, grade: false, solution: true -->
```python
def less_than(n0, n1):
    ### BEGIN SOLUTION
    return n0<n1
    ### END SOLUTION
```

```python
less_than(n[3], n[5])
```

Output:
```
True
```

```python
less_than(n[5], n[3])
```

Output:
```
False
```

```python
less_than(3, 5)
```

Output:
```
True
```

```python

```

```python

```

```python

```

```python
---
以下、採点用のセルにつき編集できない:
```

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-733c82eb66762289, points: 1, grade: true, solution: false -->
```python
assert_equal(less_than(n[3], n[5]), True)
```