```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# 近代数学とプログラミング (3)

## ベン図と論理演算

```python
n = [set()]

def succ(i):
    return {i}|n[i]

for i in range(10):
    n.append(succ(i))
n
```

Output:
```
[set(),
 {0},
 {0, 1},
 {0, 1, 2},
 {0, 1, 2, 3},
 {0, 1, 2, 3, 4},
 {0, 1, 2, 3, 4, 5},
 {0, 1, 2, 3, 4, 5, 6},
 {0, 1, 2, 3, 4, 5, 6, 7},
 {0, 1, 2, 3, 4, 5, 6, 7, 8},
 {0, 1, 2, 3, 4, 5, 6, 7, 8, 9}]
```

```python
n[5]<n[3], n[5]>n[3]
```

Output:
```
(False, True)
```

```python
(n[5]&n[3])==n[5], (n[5]&n[3])==n[3]
```

Output:
```
(False, True)
```

```python
n[5]&n[3], n[5]|n[3], n[5]^n[3], n[5]-n[3]
```

Output:
```
({0, 1, 2}, {0, 1, 2, 3, 4}, {3, 4}, {3, 4})
```

* (1834&ndash;1923) [John Venn - Wikipedia](https://en.wikipedia.org/wiki/John_Venn)
    * [Venn diagram - Wikipedia](https://en.wikipedia.org/wiki/Venn_diagram)
      * ```&``` [Intersection (set theory) - Wikipedia](https://en.wikipedia.org/wiki/Intersection_%28set_theory%29)
      * ```|``` [Union (set theory) - Wikipedia](https://en.wikipedia.org/wiki/Union_%28set_theory%29)
      * ```^``` [Symmetric difference - Wikipedia](https://en.wikipedia.org/wiki/Symmetric_difference)
      * ```-``` [Complement (set theory) - Wikipedia](https://en.wikipedia.org/wiki/Complement_%28set_theory%29)

![fig-1](https://upload.wikimedia.org/wikipedia/commons/9/99/Venn0001.svg)

![fig-2](https://upload.wikimedia.org/wikipedia/commons/3/30/Venn0111.svg)

![fig-3](https://upload.wikimedia.org/wikipedia/commons/4/46/Venn0110.svg)

![fig-4](https://upload.wikimedia.org/wikipedia/commons/5/5a/Venn0010.svg)

```python

```

```python

```

```python

```

```python
n[0], n[1]
```

Output:
```
(set(), {0})
```

```python
len(n[0]), len(n[1])
```

Output:
```
(0, 1)
```

```python
bool(0), bool(1)
```

Output:
```
(False, True)
```

```python
bool(n[0]), bool(n[1])
```

Output:
```
(False, True)
```

```python

```

```python
n[0]&n[0], n[0]|n[0], n[0]^n[0]
```

Output:
```
(set(), set(), set())
```

```python
list(map(bool, (n[0]&n[0], n[0]|n[0], n[0]^n[0])))
```

Output:
```
[False, False, False]
```

```python
n[1]&n[1], n[1]|n[1], n[1]^n[1]
```

Output:
```
({0}, {0}, set())
```

```python
list(map(bool, (n[1]&n[1], n[1]|n[1], n[1]^n[1])))
```

Output:
```
[True, True, False]
```

```python
n[0]&n[1], n[0]|n[1], n[0]^n[1]
```

Output:
```
(set(), {0}, {0})
```

```python
list(map(bool, (n[0]&n[1], n[0]|n[1], n[0]^n[1])))
```

Output:
```
[False, True, True]
```

```python

```

```python

```

```python

```

課題-3) 上で定義した自然数を表す集合のリストnの要素の大小を判定する less_than() を論理演算子``&``, ``|``, ``==``を使って定義せよ。

<!-- grade_id: cell-0d56742282275a82, points: 0, grade: false, solution: true -->
```python
def less_than(n0, n1):
    ### BEGIN SOLUTION
    return (n0&n1) == n0
    ### END SOLUTION
```

```python
less_than(n[5], n[3])
```

Output:
```
False
```

```python

```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-b9899135d48d6109, points: 1, grade: true, solution: false -->
```python
assert_equal(less_than(n[5], n[3]), False)
```