```python
# initialization for my classroom
import os
from datetime import datetime as dt

def logfile(user=os.environ.get('JUPYTERHUB_USER') or 'jovyan'):
    prefix='/srv'
    if os.path.isdir(prefix) and os.access(prefix, os.W_OK):
        prefix+=('/'+user)
        if not os.path.isdir(prefix):
            os.makedirs(prefix)
    else:
        prefix='.'
    return prefix+'/'+dt.now().strftime('%Y%m%d')+'.log'

path=logfile()
#%logstop
%logstart -otq $path append

# [python - cannot override sys.excepthook - Stack Overflow](https://stackoverflow.com/questions/1261668/cannot-override-sys-excepthook/28758396)
# https://github.com/ipython/ipython/blob/e6432249582e05f438303ce73d082a0351bb383e/IPython/core/interactiveshell.py#L1952

import sys
import traceback
import IPython

try:
    _showtraceback
except NameError:
    _showtraceback=IPython.core.interactiveshell.InteractiveShell.showtraceback

try:
    _showsyntaxerror
except NameError:
    _showsyntaxerror=IPython.core.interactiveshell.InteractiveShell.showsyntaxerror

import logging
logging.basicConfig(filename=path.replace('.log','-exc.log'), format='%(asctime)s %(message)s', level=logging.ERROR, force=True)

import sys
import traceback
import IPython

def showtraceback(self, *args, **kwargs):
    etype, value, tb = self._get_exc_info(kwargs.get('exc_tuple'))
    stb = self.InteractiveTB.structured_traceback(
        etype, value, tb, tb_offset=kwargs.get('tb_offset'))
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showtraceback(self, *args, **kwargs)

def showsyntaxerror(self, *args, **kwargs):
    etype, value, last_traceback = self._get_exc_info()
    elist = traceback.extract_tb(last_traceback) if kwargs.get('running_compiled_code') else []
    stb = self.SyntaxTB.structured_traceback(etype, value, elist)
    logging.error(os.environ.get('JUPYTERHUB_USER') or 'jovyan')
    logging.error(self.InteractiveTB.stb2text(stb))
    _showsyntaxerror(self, *args, **kwargs)

IPython.core.interactiveshell.InteractiveShell.showtraceback = showtraceback
IPython.core.interactiveshell.InteractiveShell.showsyntaxerror = showsyntaxerror
```

# 近代数学とプログラミング

## ブール代数と論理演算

* (1815&ndash;1864) [George Boole - Wikipedia](https://en.wikipedia.org/wiki/George_Boole)
  - [Boolean algebra - Wikipedia](https://en.wikipedia.org/wiki/Boolean_algebra)

```python
type(False), type(True)
```

```python
bool(0), bool(1)
```

```python
int(False), int(True)
```

```python

```

```python

```

``0`` 以外にもBoole値が ``False`` になるオブジェクト (これら以外は``True``):

```python
bool(0), bool(.0), bool(int(0)), bool(float(0))
```

```python
bool(''), bool(()), bool([]), bool({}), bool(set()), bool(range(0))
```

```python
if not '':
    print('empty')
```

```python
if not ():
    print('empty')
```

```python
if not []:
    print('empty')
```

```python
if not {}:
    print('empty')
```

```python
if not set():
    print('empty')
```

```python

```

```python

```

```python

```

### 演算を保つ写像
* [Homomorphism - Wikipedia](https://en.wikipedia.org/wiki/Homomorphism)

```python
n0 = set()
n1 = set({0})
n0, n1
```

```python
type(n0), type(n1)
```

```python
bool(n0), bool(n1)
```

```python
import itertools

list(itertools.product([n0,n1],repeat=2))
```

```python
for i,j in itertools.product([n0,n1],repeat=2):
    print(i&j, bool(i&j), bool(i)&bool(j))
```

```python
for i,j in itertools.product([n0,n1],repeat=2):
    print(i|j, bool(i|j), bool(i)|bool(j))
```

```python
for i,j in itertools.product([n0,n1],repeat=2):
    print(i^j, bool(i^j), bool(i)^bool(j))
```

```python

```

```python

```

```python

```

ブール値から整数への変換は可能だが、ブール値に変換できる対象全てを整数に変換できるわけではない

```python
n0, n1
```

```python
bool(n0), bool(n1)
```

```python
int(bool(n0)), int(bool(n1))
```

```python
int(n0), int(n1)
```

```python

```

```python

```

```python

```

課題-4) 引数が空か否かを判定する関数 ``is_not_empty`` を定義せよ。ただし、空以外の時にBoole値 ``True`` を返すこと。

<!-- grade_id: cell-c60238923b4d98f0, points: 0, grade: false, solution: true -->
```python
def is_not_empty(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
list(map(is_not_empty, ('', {}, set(), [])))
```

```python

```

```python

```

---
以下、採点用のセルにつき編集できない:

```python
from nose.tools import assert_equal, assert_true
```

<!-- grade_id: cell-30244504f92d3607, points: 1, grade: true, solution: false -->
```python
assert_equal(is_not_empty(''), False)
```